evaluation_interval: 1m

rule_files:
  - prometheus.component_alerts.yaml

tests:
  # ----- Component Deletion Alerting Tests -----
  - interval: 1m
    input_series:

      # HAS is experiencing no component deletions and no failures, should not alert
      - series: 'has_component_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+0x64'
      - series: 'has_component_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+0x64'


  - interval: 1m
    input_series:

      # HAS is experiencing component deletion failure rate of 100%, so it will be alerted
      - series: 'has_component_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'
      - series: 'has_component_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentDeletionErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing component deletion failures of 100%
              description: >-
                Component controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully delete at least 95% of components over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/has/component-delete-failed.md

  - interval: 1m
    input_series:

      # HAS experienced ~10% (11.67%) deletion failure rate, should still alert
      - series: 'has_component_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '10+10x64'
      - series: 'has_component_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x20 0+10x7 70x37'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentDeletionErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing component deletion failures of 11.67%
              description: >-
                Component controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully delete at least 95% of components over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/has/component-delete-failed.md

  - interval: 1m
    input_series:

      # HAS experienced ~10% (10.83%) deletion failure rate with counter reset, should still alert
      - series: 'has_component_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '10+10x64'
      # Increment failures by 1 for 20min, reset counter to 0 and increment failures by 10 for another 5 mins
      - series: 'has_component_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x20 0+10x5 50x38'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentDeletionErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing component deletion failures of 10.83%
              description: >-
                Component controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully delete at least 95% of components over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/has/component-delete-failed.md

  - interval: 1m
    input_series:

      # HAS experienced 5% deletion failure rate, should not alert
      - series: 'has_component_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '20+20x64'
      - series: 'has_component_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentDeletionErrors


  - interval: 1m
    input_series:

      # HAS experienced ~1% deletion failure rate, should not alert
      - series: 'has_component_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+10x64'
      - series: 'has_component_failed_deletion_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x9 5x54'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentDeletionErrors

  # ----- Component Creation Alerting Tests -----
  - interval: 1m
    input_series:

      # HAS is experiencing no component creation and no failures, should not alert
      - series: 'has_component_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+0x64'
      - series: 'has_component_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+0x64'


  - interval: 1m
    input_series:

      # HAS is experiencing component creation failure rate of 100%, so it will be alerted
      - series: 'has_component_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'
      - series: 'has_component_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentCreationErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing component creation failures of 100%
              description: >-
                Component controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully create at least 95% of components over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/has/component-create-failed.md

  - interval: 1m
    input_series:

      # HAS experienced ~10% (11.67%) creation failure rate, should still alert
      - series: 'has_component_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '10+10x64'
      - series: 'has_component_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x20 0+10x7 70x37'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentCreationErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing component creation failures of 11.67%
              description: >-
                Component controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully create at least 95% of components over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/has/component-create-failed.md

  - interval: 1m
    input_series:

      # HAS experienced ~10% (10.83%) creation failure rate with counter reset, should still alert
      - series: 'has_component_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '10+10x64'
      # Increment failures by 1 for 20min, reset counter to 0 and increment failures by 10 for another 5 mins
      - series: 'has_component_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x20 0+10x5 50x38'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentCreationErrors
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: application-service
              pod: has
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                HAS is experiencing component creation failures of 10.83%
              description: >-
                Component controller in Pod has for namespace
                application-service on cluster cluster01 is failing to
                successfully create at least 95% of components over the past hour
              alert_team_handle: <!subteam^S04MSCVRF4Z>
              team: has
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/has/component-create-failed.md

  - interval: 1m
    input_series:

      # HAS experienced 5% creation failure rate, should not alert
      - series: 'has_component_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '20+20x64'
      - series: 'has_component_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentCreationErrors


  - interval: 1m
    input_series:

      # HAS experienced ~1% creation failure rate, should not alert
      - series: 'has_component_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0+10x64'
      - series: 'has_component_failed_creation_total{namespace="application-service", pod="has", source_cluster="cluster01"}'
        values: '0x9 5x54'

    alert_rule_test:
      - eval_time: 65m
        alertname: ComponentCreationErrors
