evaluation_interval: 1m

rule_files:
  - 'prometheus.latency_pac_unprovision_alerts.yaml'

tests:
  # Scenario where only one cluster cross the 5% threshold
  # Alert triggered for one cluster
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (crosses the 5% threshold)
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="20", source_cluster="cluster01"}'
        values: '0+5x10'  # 5 requests took less than 20s
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+100x10'  # 100 total occurrences in cluster01

      # Simulating data from Cluster 2 (does not cross the 5% threshold)
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="20", source_cluster="cluster02"}'
        values: '0+95x10'  # 95 requests took less than 20s
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+100x10'  # 100 total occurrences in cluster02

    alert_rule_test:
      - eval_time: 10m
        alertname: LatencyPacUnprovision
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster01
            exp_annotations:
              summary: Latency of component PaC unprovision time exceeded
              description: >
                Time taken from PaC unprovision request till Component is unprovisioned for PaC builds has been over
                20s for more than 5% of requests during the last 5 minutes on cluster
                cluster01
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/build-service/latency_pac_unprovision.md

  # Scenario where both clusters cross the 5% threshold
  # Alert triggered for both clusters
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (crosses the 5% threshold)
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="20", source_cluster="cluster01"}'
        values: '0+2x10'  # 2 requests took less than 20s
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+100x10'  # 100 total occurrences in cluster01

      # Simulating data from Cluster 2 (also crosses the 5% threshold)
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="20", source_cluster="cluster02"}'
        values: '0+2x10'  # 2 requests took less than 20s
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+100x10'  # 100 total occurrences in cluster02

    alert_rule_test:
      - eval_time: 10m
        alertname: LatencyPacUnprovision
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster01
            exp_annotations:
              summary: Latency of component PaC unprovision time exceeded
              description: >
                Time taken from PaC unprovision request till Component is unprovisioned for PaC builds has been over
                20s for more than 5% of requests during the last 5 minutes on cluster
                cluster01
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/build-service/latency_pac_unprovision.md
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster02
            exp_annotations:
              summary: Latency of component PaC unprovision time exceeded
              description: >
                Time taken from PaC unprovision request till Component is unprovisioned for PaC builds has been over
                20s for more than 5% of requests during the last 5 minutes on cluster
                cluster02
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/build-service/latency_pac_unprovision.md

  # Scenario where neither cluster crosses the 5% threshold
  # Alert not triggered
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (does not cross the 5% threshold)
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="20", source_cluster="cluster01"}'
        values: '0+98x10'  # 98 requests took less than 20s
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+100x10'  # 100 total occurrences in cluster01

      # Simulating data from Cluster 2 (does not cross the 5% threshold)
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="20", source_cluster="cluster02"}'
        values: '0+98x10'  # 98 requests took less than 20s
      - series: 'redhat_appstudio_buildservice_PaC_unconfiguration_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+100x10'  # 100 total occurrences in cluster02

    alert_rule_test:
      - eval_time: 10m
        alertname: LatencyPacUnprovision
        exp_alerts: []  # No alerts are expected in this scenario
