evaluation_interval: 1m

rule_files:
  - 'prometheus.latency_component_onboarding_alerts.yaml'

tests:
  # Scenario where only one cluster cross the 10% threshold
  # Alert triggered for one cluster
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (crosses the 10% threshold)
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="60", source_cluster="cluster01"}'
        values: '0+10x10'  # 10 requests took less than 60s
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+100x10'  # 100 total occurrences in cluster01

      # Simulating data from Cluster 2 (does not cross the 10% threshold)
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="60", source_cluster="cluster02"}'
        values: '0+90x10'  # 90 requests took less than 60s
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+100x10'  # 100 total occurrences in cluster02

    alert_rule_test:
      - eval_time: 10m
        alertname: LatencyComponentOnboarding
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster01
            exp_annotations:
              summary: Latency of component onboarding time exceeded
              description: >
                Time spent from the moment of Component creation till the initial simple build is submitted,
                if PaC provision is requested upon the Component creation, then till the provision finishes has been over
                60s for more than 10% of requests during the last 5 minutes on cluster
                cluster01
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/build-service/latency_component_onboarding.md

  # Scenario where both clusters cross the 10% threshold
  # Alert triggered for both clusters
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (crosses the 10% threshold)
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="60", source_cluster="cluster01"}'
        values: '0+5x10'  # 5 requests took less than 60s
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+100x10'  # 100 total occurrences in cluster01

      # Simulating data from Cluster 2 (also crosses the 10% threshold)
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="60", source_cluster="cluster02"}'
        values: '0+5x10'  # 5 requests took less than 60s
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+100x10'  # 100 total occurrences in cluster02

    alert_rule_test:
      - eval_time: 10m
        alertname: LatencyComponentOnboarding
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster01
            exp_annotations:
              summary: Latency of component onboarding time exceeded
              description: >
                Time spent from the moment of Component creation till the initial simple build is submitted,
                if PaC provision is requested upon the Component creation, then till the provision finishes has been over
                60s for more than 10% of requests during the last 5 minutes on cluster
                cluster01
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/build-service/latency_component_onboarding.md
          - exp_labels:
              severity: critical
              slo: "true"
              source_cluster: cluster02
            exp_annotations:
              summary: Latency of component onboarding time exceeded
              description: >
                Time spent from the moment of Component creation till the initial simple build is submitted,
                if PaC provision is requested upon the Component creation, then till the provision finishes has been over
                60s for more than 10% of requests during the last 5 minutes on cluster
                cluster02
              alert_team_handle: <!subteam^S0500D143S8>
              team: build
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/build-service/latency_component_onboarding.md

  # Scenario where neither cluster crosses the 10% threshold
  # Alert not triggered
  - interval: 1m
    input_series:
      # Simulating data from Cluster 1 (does not cross the 10% threshold)
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="60", source_cluster="cluster01"}'
        values: '0+95x10'  # 95 requests took less than 60s
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="+Inf", source_cluster="cluster01"}'
        values: '0+100x10'  # 100 total occurrences in cluster01

      # Simulating data from Cluster 2 (does not cross the 10% threshold)
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="60", source_cluster="cluster02"}'
        values: '0+95x10'  # 95 requests took less than 60s
      - series: 'redhat_appstudio_buildservice_component_onboarding_time_bucket{le="+Inf", source_cluster="cluster02"}'
        values: '0+100x10'  # 100 total occurrences in cluster02

    alert_rule_test:
      - eval_time: 10m
        alertname: LatencyComponentOnboarding
        exp_alerts: []  # No alerts are expected in this scenario
