evaluation_interval: 1m

rule_files:
  - prometheus.oauth_alerts.yaml

tests:
  - interval: 1m
    input_series:

      # OAuth times average ~50 secs, so it will be alerted
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_sum{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '5+50x64'
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_count{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: OAuthLoginTime
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: spi-system
              container: manager
              type: GitHub
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                OAuth login average time is more than 30 sec on cluster01
              description: >-
                The average OAuth login time on cluster cluster01 has 50 sec for the last 5 minutes
              alert_team_handle: <!subteam^SLRSHSU1K>
              team: spi
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/spi/oauth_login.md

  - interval: 1m
    input_series:

      # Separate clusters reported independently
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_sum{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '5+50x64'
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_sum{source_cluster="cluster02", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '5+60x64'
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_count{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '1+1x64'
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_count{source_cluster="cluster02", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: OAuthLoginTime
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: spi-system
              container: manager
              type: GitHub
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                OAuth login average time is more than 30 sec on cluster01
              description: >-
                The average OAuth login time on cluster cluster01 has 50 sec for the last 5 minutes
              alert_team_handle: <!subteam^SLRSHSU1K>
              team: spi
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/spi/oauth_login.md
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: spi-system
              container: manager
              type: GitHub
              source_cluster: cluster02
            exp_annotations:
              summary: >-
                OAuth login average time is more than 30 sec on cluster02
              description: >-
                The average OAuth login time on cluster cluster02 has 60 sec for the last 5 minutes
              alert_team_handle: <!subteam^SLRSHSU1K>
              team: spi
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/spi/oauth_login.md


  - interval: 1m
    input_series:

      # OAuth times average 20 secs, so it will not be alerted
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_sum{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '5+20x64'
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_count{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="GitHub"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: OAuthLoginTime
        exp_alerts: []

  - interval: 1m
    input_series:

      # Non-OAuth provider must not be reported
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_sum{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="Quay"}'
        values: '5+50x64'
      - series: 'redhat_appstudio_spi_oauth_flow_complete_time_seconds_count{source_cluster="cluster01", namespace="spi-system", pod="spi-controller-manager", container="manager", type="Quay"}'
        values: '1+1x64'

    alert_rule_test:
      - eval_time: 65m
        alertname: OAuthLoginTime
        exp_alerts: []
