rule_files:
  - prometheus.release_service_alerts.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'release_processing_duration_seconds_bucket{le="3600", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_processing_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
      - series: 'release_processing_duration_seconds_bucket{le="3600", job="release", namespace="foo", source_cluster="cluster02"}'
        values: '1+8x59'
      - series: 'release_processing_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster02"}'
        values: '1+9x59'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceProcessingDuration
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              job: "release"
            exp_annotations:
              summary: >-
                90% of Releases must be processed under one hour
              description: >-
                Release service is failing to successfully process within the period of one hour for 90% of releases
              alert_team_handle: <!subteam^S03SVBS426R>
              team: release

  - interval: 1m
    input_series:
      - series: 'release_pre_processing_duration_seconds_bucket{le="10", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_pre_processing_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
      - series: 'release_pre_processing_duration_seconds_bucket{le="10", job="release", namespace="foo", source_cluster="cluster02"}'
        values: '1+8x59'
      - series: 'release_pre_processing_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster02"}'
        values: '1+9x59'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServicePreProcessingDurationSeconds
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              job: "release"
            exp_annotations:
              summary: >-
                90% of Releases must start processing under 10 seconds
              description: >-
                Release service is failing to start processing under 10 seconds for 90% of releases
              alert_team_handle: <!subteam^S03SVBS426R>
              team: release

  - interval: 1m
    input_series:
      - series: 'release_validation_duration_seconds_bucket{le="5", job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+8x59'
      - series: 'release_validation_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster01"}'
        values: '1+9x59'
      - series: 'release_validation_duration_seconds_bucket{le="5", job="release", namespace="foo", source_cluster="cluster02"}'
        values: '1+8x59'
      - series: 'release_validation_duration_seconds_count{job="release", namespace="foo", source_cluster="cluster02"}'
        values: '1+9x59'
    alert_rule_test:
      - eval_time: 1h
        alertname: ReleaseServiceValidationDurationSeconds
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              job: "release"
            exp_annotations:
              summary: >-
                90% of Releases must be validated under 5 seconds
              description: >-
                Release service is failing to run the validations under 5 seconds for 90% of releases
              alert_team_handle: <!subteam^S03SVBS426R>
              team: release
